---
title       : Hypothesis Testing
description : Introducing hypothesis testing


--- type:NormalExercise lang:r xp:100 skills:1 key:5d094e21fe
## Binomial test


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
