---
title       : Probability and Risk
description : Understanding uncertainty


--- type:VideoExercise lang:r xp:50 skills:1 key:a234a254b1
## Welcome to uncertainty


*** =video_link
//player.vimeo.com/video/154783078


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:cd4dd49ae9
## How am I going to die?

What is the risk of being killed by lightning in the next year?

*** =instructions
- 1 in 8000
- 1 in 50 000
- 1 in 500 000
- 1 in 30 000 000
- 1 in 40 000 000
- 1 in 45 000 000
*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}
test_mc(correct = 4)
```


--- type:VideoExercise lang:r xp:50 skills:1 key:94eafe5ef2
## What is probability?


*** =video_link
//player.vimeo.com/video/154783078


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:d639cec9ad
## Probability is...

Which is the best definition of probability?


*** =instructions

- The exact proportion of a small number of independent trials which give a particular outcome.
- The approximate proportion of a small number of independent trials which give a particular outcome.
- The exact proportion of a large number of independent trials which give a particular outcome.
- The approximate proportion of a large number of independent trials which give a particular outcome.
*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}
test_mc(correct = 4)
```


--- type:VideoExercise lang:r xp:50 skills:1 key:dca74ed106
## What's with the dice?


*** =video_link
//player.vimeo.com/video/154783078

