---
title       : Hypothesis Testing
description : Introducing hypothesis testing





--- type:VideoExercise lang:r xp:50 skills:1 key:b6260c0853
## Hypothesis testing


*** =video_link
//player.vimeo.com/video/227739062




--- type:PlainMultipleChoiceExercise lang:r xp:50 skills:1 key:50ed979fc5
## Hypothesis testing 2


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```



--- type:VideoExercise lang:r xp:50 skills:1 key:0c94221b8b
## Is it a fair coin?


*** =video_link
//player.vimeo.com/video/227743190



--- type:PlainMultipleChoiceExercise lang:r xp:50 skills:1 key:e5019bee9f
## Is it a fair coin? (2)


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```



--- type:VideoExercise lang:r xp:50 skills:1 key:f18d570d6a
## One tail or two?


*** =video_link
//player.vimeo.com/video/227745553



--- type:PlainMultipleChoiceExercise lang:r xp:50 skills:1 key:8276675a79
## One tail or two? (2)


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```



--- type:VideoExercise lang:r xp:50 skills:1 key:6a1acb028b
## The binomial test in R


*** =video_link
//player.vimeo.com/video/227746218

--- type:NormalExercise lang:r xp:100 skills:1 key:5d094e21fe
## The binomial test in R (2)


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```



--- type:VideoExercise lang:r xp:50 skills:1 key:18228d8d30
## What is a p value?


*** =video_link
//player.vimeo.com/video/227747288


--- type:PlainMultipleChoiceExercise lang:r xp:50 skills:1 key:e8ee19a7a2
## What is a p value? (2)


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```


--- type:VideoExercise lang:r xp:50 skills:1 key:4c7caf8521
## Possible outcomes of hypothesis tests


*** =video_link
//player.vimeo.com/video/227749035





--- type:PlainMultipleChoiceExercise lang:r xp:50 skills:1 key:742d710903
## Possible outcomes of hypothesis tests (2)


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```
--- type:VideoExercise lang:r xp:50 skills:1 key:dc0eebd374
## Power


*** =video_link
//player.vimeo.com/video/227751750


--- type:PlainMultipleChoiceExercise lang:r xp:50 skills:1 key:b06700136e
## Power (2)


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:ba89b22b24
## Key points


*** =video_link
//player.vimeo.com/video/227754182
