---
title       : Correlation and Regression
description : Investigating the association between two continuous variables



--- type:NormalExercise lang:r xp:100 skills:1 key:377c84540a
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
